
// I AM NOT DONE

// Ressources
// https://docs.cairo-lang.org/cairozero/how_cairo_works/cairo_intro.html#registers
// https://docs.cairo-lang.org/cairozero/how_cairo_works/functions.html#function-arguments-and-return-values

// TODO
// Rewrite this function body in a high level syntax
func ret_42() -> felt {
    // [ap] = 42; ap++
    // ret
}

// TODO
// Rewrite this function body in a low level syntax, using registers
func ret_0_and_1() -> (zero: felt, one: felt) {
    // return (0, 1);
}
