%lang starknet

// Starknet provides a module management system.
// It is very similar to the Python's one.

// I AM NOT DONE

// TODO: add the module imports needed to make the test pass!
from starkware.cairo.common.cairo_builtins import HashBuiltin

// You can ignore what follows for now
@external
func test_ok{pedersen_ptr : HashBuiltin*, range_check_ptr}(){
    return();
}
