// Hints

// Hints are the Cairo way to defer part of the program execution to an external program
// Memory can be accessed and assigned inside a hint by using variables identifiers.
// e.g., inside a hint variable `a` is accessed through `ids.a`

// I AM NOT DONE

// TODO: Assign the value of `res` inside a hint.

func basic_hint() -> felt{
    alloc_locals;
    local res;
    // TODO: Insert hint here
    return res;
}
